// primitive_types1.rs
//
// Fill in the rest of the line that has code missing! No hints, there's no
// tricks, just get used to typing these :)

// I AM NOT DONE

fn main() {
    // Booleans (`bool`)

    let is_morning = true;
    if is_morning {
        println!("Good morning!");
    }

    let bol = false;
    let is_evening = if bol {// Finish the rest of this line like the example! Or make it be false!
        false
    }else{
        true
    };
    if is_evening {
        println!("Good evening!");
    }
}
